Sleep need driven oscillation of glutamate synaptic phenotype

Author response:

The following is the authors’ response to the original reviews.

We greatly appreciate reviewer 2 comments with both insightful and clearly evaluated assessments of this study that include, much appreciated reframing and evaluation of the study’s advances in the sleep field. It is a constructive review and provides considerable added value to this study in better defining the biological significance of the findings, including both advances and limitations.

Reviewer 2 nicely summarized the work as “…highlight(ing) the accumulation and resolution of sleep need centered on the strength of excitatory synapses onto excitatory neurons.”. The reviewer succinctly placed one of the main electrophysiological findings in context of one of the sleep field’s most prevalent views, “that LTP associated with wake, leads to the accumulation of sleep need by increasing neuronal excitability, and by the "saturation" of LTP capacity.” It has been speculated that “This saturation subsequently impairs the capacity for further ongoing learning. This new data provides a satisfying mechanism of this saturation phenomenon (and its restoration by recovery sleep) by introducing the concept of silent synapses.” We want to emphasize that sleep need and its resolution involves more than just homeostasis of excitatory synaptic strength but may also be extended to include homeostasis of excitatory synaptic potential to undergo LTP (a homeostasis of meta-plasticity), with implications for learning and memory.

Reviewer 2 also identified another advance made by this study, summarized as, “The new snRNAseq dataset indicates the sleep need is primarily seen (at the transcriptional level) in excitatory neurons, consistent with a number of other studies.” References for these studies are nicely provided by the reviewer. Our analysis of this data extends the evidence for transcriptional sleep-need-driven changes, observed by us and others in excitatory neurons to more particularly involve the excitatory neurons in layers 2-5, targeting intra-telencephalic neurons.

Reviewer 2, importantly noted, “New snRNAseq analysis indicates that SD drives the expression of synaptic shaping components (SSCs) consistent with the excitatory synapse as a major target for the restorative basis of sleep function”, and that “SD-induced gene expression is also enriched for autism spectrum disorder (ASD) risk genes”. These comments are well appreciated as they emphasize that beyond identification of the major target cell type of sleep function, the major sleep-target, gene-ontological characteristics are starting to be addressed.

Reviewer 2 commented on the molecular sleep model, making a key observation that “SDinduced gene expression in excitatory neurons overlaps with genes regulated by the transcription factor MEF2C and HDAC4/5 (Figure 4),” and accurately discusses the significance with respect to the proposed model.

We are in complete agreement with the observation that the molecular sleep model presented is not “definitively supported by the new data and in this regard should be viewed as a perspective…”. One of the more glaring gaps in supporting evidence is the absence of understanding of the role of HDAC4/5 (part of the SIK3-HDAC4/5 pathway) in sleep need modulation of excitatory synapses. Resolution of this issue might be approached by assessment of the synaptic effects of constitutively nuclear HDAC4/5. The current study provides a first step in the assessment by showing a correlation between HDAC4/5 and MEF2c target genes and a subset of differentially expressed synaptic shaping component (SSC) genes that modulate excitatory synapse strength and phenotype. However, the functional studies have yet to be completed. Complimentary studies on SD-induced SSC-DEGs (identified in this study) are also needed for follow-up characterization of their sleep need induced functional impact (both strength and meta-plasticity modulation) on the most relevant excitatory synapses (as identified in the current study).

We agree with both reviewers 1 and 2 that, “Additional work is also needed to understand the mechanistic links between SIK3-HDAC4/5 signaling and MEF2C activity”. Reviewer 2 clarifies the key unresolved issue as, “cnHDAC4/5 suppresses NREM amount and NREM SWA but had no effect on the NREM-SWA increase following SD (Zhou et al., Nature 2022). Loss of MEF2C in CaMKII neurons had no effect on NREM amount and suppressed the increase in NREM-SWA following SD (Bjorness et al., 2020)”. One may conclude with reviewer 2, “These instances indicate that cnHDAC4/5 and loss of MEF2C do not exactly match suggesting additional factors are relevant in these phenotypes.”

An understanding of the mechanism(s) responsible for the relationship between sleep need and SWA are critical to the evaluation of sleep need’s correlation with sleep DEGs and synaptic transmission, including “additional factors” as suggested by reviewer 2. SWA might result from a decrease of cortical glutamatergic neurotransmission below some threshold, which might occur in response to prolonged waking (possibly in response to waking activity-induced local increases of adenosine?), rather than being a cause of, or, being intimately involved in resolving sleep need.

An increase of SWA in association with SD can result directly from an acute SD-induced increase in local adenosine concentration. This will elicit an ADORA1-mediated down-regulation of glutamate excitatory neurotransmission in the cortex (Bjorness et al., 2016) and in cholinergic arousal centers (Rainnie et al., 1994; Porkka-Heiskanen et al., 1997; Portas et al., 1997; Li et al., 2023). When MEF2c is derepressed by chronic loss of HDAC4 function, SWA is facilitated (Kim et al., 2022). It is plausible that loss of HDAC4 function contributes to the increased SWA by downscaling glutamate excitatory transmission (independent of sleep need). This is expected to result from derepressed, MEF2c mediated sleep-gene expression.

Similarly, over-expression of constitutively active HDAC4 (cnHD4) can contribute to chronic upscaling of cortical glutamate synaptic strength to depress SWA (again, independent of sleep need). Thus, facilitation or depression of SWA correlates with up or down scaling effects on cortical glutamate neurotransmission, respectively, even in the absence of a direct effects on sleep need (Figure 4D). Many reagents that reduce the excitability of glutamate pyramidal cells by various mechanisms, including anesthetics like isoflurane, barbiturates or benzodiazepines in addition to those activating ADORA1, increase SWA. Finally, it is important to acknowledge that direct evidence for this proposed link of SWA to cortical glutamate transmission remains in need of further investigation. Thus, SWA may reflect generalized cortical glutamate synaptic activity whether modulated by sleep function or by other agents.

Still, other factors that can have a role mediating some of the mis-match between cnHD4/5 DEGs and Mef2c-cKO DEGs, include the broader over-expression of AAV-cnHD4 compared to CamKII- driven Cre KO of Mef2c. The cnHD4 overexpression can increase arousal center activity in the hypothalamus and other arousal areas to interfere with SWA, but not to the exclusion of SD-DEG repression resulting from a repression of MEF2c-mediated sleep gene expression.

The critique by reviewer 1 raises a number of important technical issues with this study. A key, potentially critical issue raised by reviewer 1, is that of our method of experimental sleep deprivation (ESD). The reviewer suggests that “…neuronal activity/induction of plasticity”, peculiar to the ESD methodology employed in this study, “…rather than sleep/wake states are responsible for the observed results…”.

In this study, a slow-moving treadmill (SMTM; 0.1km/hour, as stated in the methods), requiring locomotion to avoid bumping into the backwall of a false bottomed plexiglass cage was used to induce ESD. A mouse, in its home cage, typically moves much faster than 0.1km/hour and the mouse is able to eat and drink freely while in the cage (see file: video 1). Furthermore, our observations using a beam-break cage, indicate that mice spontaneously travel for comparable to longer distances over 6 hours than the treadmill moves (during the ESD of 6 hours). Finally, our EEG recordings of mice on the active treadmill show 100% waking while it is on (Bjorness et al., 2009), whereas prevention of NREM sleep (including transition time) using the “gentle handling” (GH) technique occurs depending on the diligence of the experimenter.

The accommodation (one week prior to ESD) included exposure to the treadmill-on for 30minutes ~ZT=2 & ZT= 14 hours (now spelled out in the “Materials & Methods” section). Thus, the likelihood of motor learning seems vanishingly small.

As with all ESD methods, there must be some associated increase in sensory and motor neuronal activity to drive arousal and prevent transition to sleep. For example, the more widely employed GH method of ESD involves sensory stimulation (tactile and or auditory) of sufficient intensity to induce postural change from that associated with sleep to that associated with wake (often involving some locomotion). Like the SMTM, both sensory and motor systems are likely to be engaged. Unlike the SMTM method, the stimulation used in GH is variably-intermittent from mouse to mouse and from experimenter to experimenter as it is applied only when the experimenter judges the mouse to be falling asleep. . It can even be argued that the varied and unpredictable ways in which these interactions happen cause plastic changes with a higher likelihood than the constant slow motion of a treadmill – the mice know how to walk, after all. In other protocols, novel objects are introduced to the animals – those will certainly trigger plastic processes –something that is avoided using a slow-running treadmill to which the mouse has been accommodated, for sleep deprivation.

The changes induced by SMTM technique are reproducible and induce arousal by somatic stimulation of sufficient intensity to induce natural motor activity as with GH. All ESD methods induce motor activity and it is reasonable to speculate that induced, motor activity is essential for effective ESD for the prolonged durations (>4 hours in mice) that elicit high sleep need. Electrophysiological assessment of SD-evoked increases in mEPSC amplitude and frequency using GH-ESD (Liu et al., 2010) are similar in all respects to our observations of the response to SMTMESD (Bjorness et al., 2020). Further studies might directly address a comparison of SMTM-ESD to GH-ESD as suggested by reviewer 1 but are regrettably outside the scope and resources of our study.

The model presented in Figure 4C is consistent with the experimental findings with respect to the observed electrophysiological changes (including loss of silent synapses and increased AMPA/NMDA ratio after ESD of 6 hours) and altered gene expression that includes enrichment of SSC genes, many of which (7 candidates are listed) can affect both AMPA/NMDA ratio and silent synapses. No claim of mechanism linking the changed expression to altered AMPAR or NMDAR activity can be made at this point, even as to polarity of gene expression, related to electrophysiological outcome. Furthermore, some transcripts may involve receptor trafficking while others more directly affect activated receptor function. To help illustrate the complexity of interpreting gene up-regulation, consider the following hypothetical scenario. If a gene like upregulated Grin3a acts rapidly, it may facilitate reduction of NMDAR function (decreasing plasticity) during ESD, whereas upregulation of a gene like Kif17, if acting in a more delayed manner, might enhance NMDAR surface expression and activity (increasing silent synapses) in response to ESD, during recovery sleep. Relevant references, consistent with these various outcomes are supplied in the manuscript but further investigation is clearly needed, or as reviewer 2 so aptly commented, this work “…provides a framework to stimulate further research and advances on the molecular basis of sleep function”.

Several issues are raised by reviewer 1 concerning the electrophysiological methodology and statistical assessment. In regard to the former, we closely followed established protocols employed in the frontal neocortex (Myme et al., 2003). We did not include the details for series resistance monitoring. Series resistance values ranged between 8 and 15 MOhm and experiments with changes larger than 25% not used for further analyses. Thank you for bringing this oversight on our part, to our attention. This essential information, that is unfailingly gathered for all our whole cell recordings, is now added to the version of record.

The -90 mV holding potential was chosen according to precedent (Myme et al., 2003). It increases driving force and permits lower stimulus strength for the same response size – reducing the likelihood for polysynaptic responses. Experiments with multiple response peaks at -90 mV were not included in the analysis. The -90 mV holding potential also increases NMDA receptor Mg++ block resulting in a minimally contaminated AMPA response. This information is now added to our submitted version of record.

The statistical assessments shown in Table 1 refer to two sets of data measured from 3X2=6 different cohorts for each sleep condition (CS, SD, RS): 1) AMPA & NMDA EPSCs and 2) AMPA/NMDA FR ratios (FRR; now bolded in row 1, second tab, Table S1). As stated in the results section, “A two-way ANOVA analysis showed a significant interaction between AMPA matched to NMDA EPSC response for each neuron, and sleep condition (F (2, 21) = 7.268, p<0.004; Figure 1 A, C, E). When considered independently, neither the effect of sleep condition nor of EPSC subtype reached significance at p<0.05 (Figure 1 C)”.

As noted by reviewer 1, we inadvertently dropped one of the data points from the RS FR and FR ratio (FRR) statistical analysis (raw data in the third tab of Table S1, statistical data in fourth and fifth tab and illustrated in figure 1 F). Thanks to this appreciated, rigorous review, we can correct the oversight (using raw data unchanged in Table S1, third tab). The Table S1 and figure 1 F are now corrected for the version of record. For better clarity, we now use two tabs, the fourth and fifth tabs, respectively of Table S1, for separate stat analyses of FR and FRR data.

The significance of the AMPA/NMDA FRR across sleep conditions was assessed with the KruskalWallis test, a non-parametric method. The two-stage linear step-up procedure of Benjamini, Krieger, and Yekutieli (BKY) was used to control for the FDR across multiple sleep conditions, in the non-parametric Kruskal-Wallis test but it is usually less powerful than tests presuming normal distributions like the one-way ANOVA and Holm-Sidak’s test. We have now added re-analyzed FRR across CS, SD and RS conditions using a normal one-way ANOVA (Table S1, tab5). The results now read, “The difference between sleep conditions and FRR is significant (F (2, 19) = 11.3, Table S1, tab5). Multiple comparisons (Holm-Sidak, Table S1, tab5) indicate the near absence of silent synapses was reversed by either CS or RS (SD/CS; p<0.0011 and SD/RS: p<0.0006; Table S1, tab 5; Figure 1 F).”. These analyses compare well to the non-parametric assessment using the KruskalWallis test (significant at p= 0.0006) with BYK correction for multiple comparison analysis to give for CS-SD, p<= 0.0262 and for RS-SD, p<= 0.0006 (statistics also shown in Table S1, tab5). [Also shown in tab5 is the “standard approach of correcting for family wise error rate”, namely, Dunn’s test. It is more conservative but less powerful than the BYK correction- in general the tradeoff of greater power/ less conservative is better tolerated when many comparisons are made, however, it can be argued that in the present analysis type 2 errors are also potentially misleading and thus not well tolerated.] The modifications of our statistical analyses, inspired by reviewer 1, did not affect the interpretation of the data nor the conclusions.

Bjorness TE, Kelly CL, Gao T, Poffenberger V, Greene RW (2009) Control and function of the homeostatic sleep response by adenosine A1 receptors. The Journal of neuroscience : the official journal of the Society for Neuroscience 29:1267-1276.

Bjorness TE, Dale N, Mettlach G, Sonneborn A, Sahin B, Fienberg AA, Yanagisawa M, Bibb JA, Greene RW (2016) An Adenosine-Mediated Glial-Neuronal Circuit for

Homeostatic Sleep. The Journal of neuroscience : the official journal of the Society for Neuroscience 36:3709-3721.

Bjorness TE, Kulkarni A, Rybalchenko V, Suzuki A, Bridges C, Harrington AJ, Cowan CW, Takahashi JS, Konopka G, Greene RW (2020) An essential role for MEF2C in the cortical response to loss of sleep in mice. Elife 9.

Kim SJ et al. (2022) Kinase signalling in excitatory neurons regulates sleep quantity and depth. Nature 612:512-518.

Li B, Ma C, Huang YA, Ding X, Silverman D, Chen C, Darmohray D, Lu L, Liu S, Montaldo G, Urban A, Dan Y (2023) Circuit mechanism for suppression of frontal cortical ignition during NREM sleep. Cell 186:5739-5750 e5717.

Liu ZW, Faraguna U, Cirelli C, Tononi G, Gao XB (2010) Direct evidence for wake-related increases and sleep-related decreases in synaptic strength in rodent cortex. The Journal of neuroscience : the official journal of the Society for Neuroscience 30:8671-8675.

Myme CI, Sugino K, Turrigiano GG, Nelson SB (2003) The NMDA-to-AMPA ratio at synapses onto layer 2/3 pyramidal neurons is conserved across prefrontal and visual cortices. Journal of neurophysiology 90:771-779.

Porkka-Heiskanen T, Strecker RE, Thakkar M, Bjorkum AA, Greene RW, McCarley RW (1997) Adenosine: a mediator of the sleep-inducing effects of prolonged wakefulness. Science 276:1265-1268.

Portas CM, Thakkar M, Rainnie DG, Greene RW, McCarley RW (1997) Role of adenosine in behavioral state modulation: a microdialysis study in the freely moving cat. Neuroscience 79:225-235.

Rainnie DG, Grunze HC, McCarley RW, Greene RW (1994) Adenosine inhibition of mesopontine cholinergic neurons: implications for EEG arousal. Science 263:689692.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Sleep loss increases AMPA-synaptic strength and number in the neocortex. However, this is only part of the synaptic sleep loss response. We report increased AMPA/NMDA EPSC ratio in frontal-cortical pyramidal neurons of layers 2-3. Silent synapses are absent, decreasing the plastic potential to convert silent NMDA to active AMPA synapses. These sleep loss changes are recovered by sleep. Sleep genes are enriched for synaptic shaping cellular components controlling glutamate synapse phenotype, overlap with autism risk genes and are primarily observed in excitatory pyramidal neurons projecting intra-telencephalically. These genes are enriched with genes controlled by the transcription factor, MEF2c and its repressor, HDAC4. Thus, sleep genes under the influence of MEF2c and HDAC4, can provide a framework within which motor learning and training occurs mediated by sleep-dependent oscillation of glutamate-synaptic phenotypes.

Revised version: Sleep loss increases AMPA-synaptic strength and number in the neocortex. However, this is only part of the synaptic sleep loss response. We report increased AMPA/NMDA EPSC ratio in frontal-cortical pyramidal neurons of layers 2-3. Silent synapses are absent, decreasing the plastic potential to convert silent NMDA to active AMPA synapses. These sleep loss changes are recovered by sleep. Sleep genes are enriched for synaptic shaping cellular components controlling glutamate synapse phenotype, overlap with autism risk genes and are primarily observed in excitatory pyramidal neurons projecting intra-telencephalically. These genes are enriched with genes controlled by the transcription factor, MEF2c and its repressor, HDAC4. Sleep genes can thus provide a framework within which motor learning and training occurs mediated by sleep-dependent oscillation of glutamate-synaptic phenotypes.

Added in the revised version: Main Text

Removed in the revised version: Introduction

CNS control of arousal maintains waking neuronal activities needed for foraging, danger avoidance and reproductive behaviors. Reduction of arousal to unconsciousness can gate sleep, however, the function served by this sleep, in contrast to wake function, remains enigmatic. A reasonable teleologically oriented goal is more effective future foraging, danger avoidance and reproductive behaviors, based on procedural and episodic learning from past waking experience. Effective learning involves memory of waking experience and underlying long-term plasticity (LTP) (Moser et al., 1998; Morris et al., 2003; Stickgold, 2005). At the cellular level, extensive activity dependent LTP of cortical glutamate synapses can lead to a well-documented increase of synaptic strength of cortical glutamatergic synapses during waking. It is reflected by increased frequency and amplitude of glutamatergic miniature synaptic currents (mEPSC) that is restored by subsequent sleep (Liu et al., 2010; Torrado Pacheco et al., 2021), consistent with the synaptic homeostasis hypothesis of sleep (SHY) (Tononi and Cirelli, 2014).

Our recent findings indicate that during waking, not only are individual synapses strengthened (overall glutamate mEPSC amplitude in layer 2-3 pyramidal cells is increased by SD) but, also, the number of electro-physiologically active synapses is increased (increased frequency in correlation with unaltered or decreased probability of release) (Bjorness et al., 2020). The molecular underpinnings of these changes in glutamate synaptic strength and number were examined to determine the changes in the glutamate synaptic phenotype and what the implication(s) of these changes, if any, might be with regards to propensity for synaptic plasticity (meta-plasticity).

We investigated sleep/wake mediated electrophysiological changes in glutamatergic synapses in motor cortex, layer 2-3, pyramidal neurons since our previous work demonstrates wake/sleep up-regulation and down-regulation, respectively, in these cells’ glutamate synapses (Bjorness et al., 2020). We observed a wake/sleep oscillation of AMPA/NMDA synaptic response ratio in pyramidal neurons of motor neocortex coupled with an oscillating fraction of silent synapses (NMDAR-mediated synaptic response but no AMPAR mediated response) in response to sleep need. We then, extended these physiological findings to include characterization of sleep-need related differential expression of glutamatergic synaptic shaping genes that can control the observed oscillating glutamate synapse phenotype.

Results

Sleep need driven, glutamate, synaptic phenotype

Three cohorts of C57 BL/6 male mice were allowed either to 1) sleep ad lib (control sleep, CS); 2) prevented from sleeping from ZT=0 to 6h (sleep deprived, SD); or, 3) after 4h SD allowed recovery sleep, ad lib for 2h (RS). Ex Vivo acute brain slices from the motor cortex were then prepared from these animals at ZT=6h and whole cell recordings obtained from layer 2/3 pyramidal neurons (see methods for details). Locally evoked glutamatergic EPSCs recorded at Vhold = -90mV had little NMDAR contribution due to Mg++ block. These were compared in the same cell, to EPSCs evoked at Vhold = +50mV with little AMPAR contribution due to reduced AMPAR driving force, but, having large NMDAR currents due to voltage sensitive relief from Mg++ blockade (NMDAR currents had reversed polarity since ENMDAR =∼0mV). For each recorded neuron (one neuron/ slice/ animal) an average response to repeated stimuli (>10 stimuli, 100pA, applied every 10sec) to elicit a EPSC was determined at each holding potential. A two-way ANOVA analysis showed a significant interaction between AMPA matched to NMDA EPSC response for each neuron, and sleep condition (F (2, 21) = 7.268, p<0.004; Figure 1 A, C, E). When considered independently, neither the effect of sleep condition nor of EPSC subtype reached significance at p<0.05 (Figure 1 C).

Sleep need-dependent responses of AMPA/NMDA ratio and silent synapses in motor cortex. A) Examples of AMPA currents at -90 mV holding potential and NMDA currents at +50 mV holding potential are shown for: control sleep (CS), 6 h of sleep deprivation (SD) and 4 h of sleep deprivation followed by 2 h of recovery sleep (RS). Traces (100msec duration) are scaled to the NMDA current for comparison (NMDA current measured @ 40msec after AMPA peak current). B) Examples of successes (black) and failures (red) at -90 mV (left, AMPA) and +50 mV (right, NMDA) after minimal stimulation of excitatory inputs to motor cortex pyramidal neurons are shown; Top row: CS sleep, middle row: SD and bottom row: RS (Cal. 10pAX20msec). Rate of failures (% of all stimuli delivered) for AMPA EPSCs (blue) and NMDA EPSCs (red) in the three conditions (+/- sem). C), D) Average (+/- sem) AMPAR and NMDAR EPSC responses (unmatched) and failure rates, respectively, for each sleep condition. E), F) Matched AMPA/NMDA EPSC response and AMPA/NMDA failure rate ratios, respectively, are shown for the three conditions (N=1 cell/slice, 2-3 slice/animal, 3 animals/condition).

Earlier version: To control for neuron to neuron EPSC variability (due to variability of stimulation-evoked afferent activation), we calculated the AMPA/NMDA ratio, matched for each recorded neuron. This showed a significantly greater AMPA/NMDA EPSC ratio comparing SD to CS or RS (p<=0.03 for CS, p<=0.001 for RS; Kruskal-Wallis test and two-stage linear step-up procedure of Benjamini, Krieger, Yekutieli; K-W, BKY; Table 1).

Revised version: To control for neuron-to-neuron EPSC variability (due to variability of stimulation-evoked afferent activation), we calculated the AMPA/NMDA ratio, matched for each recorded neuron. This showed a significantly greater AMPA/NMDA EPSC ratio (F (2, 21)= 3.203) and comparing SD to CS or RS, p<=0.03 or p<=0.02, respectively; Dunnett’s multiple comparisons test; Table S1, tab 2).

Earlier version: Many glutamatergic synapses on cortical pyramidal cells do not respond to activation at resting membrane potential, yet they do have an NMDAR synaptic component that can be observed at Vhold=+50mV. After activation by LTP-evoking stimulation, these ‘silent’ synapses can convert to active, non-silent synapses by virtue of LTP-induced AMPAR insertion into the post-synaptic active zone (Liao et al., 1995). The ratio of AMPAR component failure rate (FR; Vhold= -90mv; FR-90 for active synapses, As) matched with the same neuron’s FR+50 (Vhold= +50mV; for As + silent synapses, Ss) has been used to estimate the change in the fraction of Ss/As in response to LTP(Liao et al., 1995). We employed a similar analysis to examine the effects of SD and sleep on functional synaptic AMPAR and NMDAR composition.

Revised version: Many glutamatergic synapses on cortical pyramidal cells do not respond to activation at resting membrane potential, yet they do have an NMDAR synaptic component that can be observed at Vhold=+50mV. After activation by LTP-evoking stimulation, these ‘silent’ synapses can convert to active, non-silent synapses by virtue of LTP-induced AMPAR insertion into the post-synaptic active zone (Liao et al., 1995). The ratio of AMPAR component failure rate (Vhold= -90mv; FR-90 for active synapses, As) matched with the same neuron’s FR+50 (Vhold= +50mV; for As + silent synapses, Ss) has been used to estimate the change in the fraction of Ss/As in response to LTP (Liao et al., 1995). We employed a similar analysis to examine the effects of SD and sleep on functional synaptic AMPAR and NMDAR composition. We observed a significant interaction between FR’s and sleep conditions, assessed by a two-way ANOVA mixed effects model (F (2, 19) = 5.550, p = 0.0126; Figure 1 B1,2,3 D F; Table S1, tab 4). Multiple comparisons revealed no significant difference between FR-90 and FR+50 in SD indicating far fewer numbers of silent synapses (Ss) in this condition. However, both CS and RS restored the relatively decreased number of Ss compared to SD. FR-90 is significantly smaller than FR+50 for both CS and RS (p<0.0001 and <0.0005, respectively).

Removed in the revised version: We observed a significant interaction between FR’s and sleep conditions, assessed by a two-way ANOVA (F (2, 18) = 5.433 p=0.0143; Figure 1 B1,2,3 D F; Table 1). Multiple comparisons (K-Y BKY test) revealed no significant difference between FR-90 and FR+50 in SD indicating far fewer numbers of silent synapses (Ss) in this condition. However, both CS and RS restored the relatively decreased number of Ss compared to SD. FR-90 is significantly smaller than FR+50 for both CS and RS (p<0.0001 and <0.0007, respectively).

Earlier version: The condition specific ratio of FR+50mV to FR-90mV is a power function of the probability of failure for each of the stimulated synapses (1-prelease)stimulated-synapse and thus, inversely proportional (since (1-prelease)<1) to the number of silent synapses (Ss). A binomial probability distribution model is employed with the assumptions that: 1) each stimulus activates the same set of pre-synaptic terminals for a given neuron; 2) each terminal has a constant prelease; 3) P(failure) is experimentally estimated by the FR (number of null synaptic responses to a set of stimuli / total number of stimuli). For a given number of stimulus-activated terminals, FR = (1-prelease)#activated terminals. The FR+50mV/FR-90mV = (1-prelease)As+Ss/(1-prelease)As and ln(FR+50mV)/ln(FR-90mV) =(As+Ss)/As=1+Ss/As. SD’s FR+50mV/FR-90mV=∼1, indicating ∼no silent synapses. The ratio of Ss/As for CS=1.44 and for RS=1.46, compared to SD=0.09. Thus, the near absence of silent synapses was reversed by either CS or RS (K-W BKY test; p<0.01 comparing ratios SD/CS, p<0.001 for SD/RS; Table 1; Figure 1 F). This suggests there are few if any silent synapses available for conversion by LTP to active synapses following SD and this availability is recovered by sleep.

Revised version: The condition specific ratio of FR+50mV to FR-90mV (FRR) is a power function of the probability of failure for each of the stimulated synapses (1-prelease)stimulated-synapse and thus, inversely proportional (since (1-prelease)<1) to the number of silent synapses (Ss). A binomial probability distribution model is employed with the assumptions that: 1) each stimulus activates the same set of pre-synaptic terminals for a given neuron; 2) each terminal has a constant prelease; 3) P(failure) for a given set of stimulus-activated terminals = (1-prelease)#activated terminals. The FR+50mV/FR-90mV = (1-prelease)As+Ss/(1-prelease)As and ln(FR+50mV)/ln(FR-90mV) =(As+Ss)/As = 1+Ss/As. SD’s FR+50mV/FR-90mV=∼1, indicating ∼no silent synapses. The FRR of Ss/As for CS=1.44 and for RS=1.46, compared to SD=0.09. The difference between sleep conditions and FRR is significant (F (2, 19) = 11.3, Table S1, tab5). Multiple comparisons (Holm-Sidak, Table S1, tab5) indicate the near absence of silent synapses was reversed by either CS or RS (SDxCS; p<0.0011 and SDxRS: p<0.0006; Table S1, tab 5; Figure 1 F). This suggests there are few if any silent synapses available for conversion by LTP to active synapses following SD and this availability is recovered by sleep.

The sleep-dependent electrophysiological findings raise several questions, including: 1) are other cell types besides the layer 2-3 pyramidal neurons from which we recorded, involved; 2) can changes in the transcriptome account, at least in part, for sleep-related synaptic changes in function, and if so, what are the specific functions of those genes and 3) what is the upstream sleep-dependent control of changes in gene expression?

Cell type/subtype specific gene expression in response to sleep loss

To examine the molecular mechanisms responsible for an SD response, we characterized single nuclei transcriptomes of cells from motor cortex of mice from two cohorts, an ad lib or control sleep (CS) group and an SD group (sleep deprived from ZT= 0-6hr; see detailed methods).

Added in the revised version: Transcript libraries from each of the two cohorts (CS, n=4 and SD, n=4), were prepared from single nuclei isolated from mouse frontal cortex. Nuclei from each cohort were clustered and visualized using Uniform Manifold

Earlier version: Transcript libraries from each of the two cohorts (CS, n=4 and SD, n=4), were prepared from single nuclei isolated from mouse frontal cortex. Nuclei from each cohort were clustered and visualized using Uniform Manifold Approximation and Projection (UMAP) and annotated using multi-modal characteristics from the Brain Initiative Cell Census(Network, 2021) (Fig. 2 A,B). The sleep condition had no significant effect on the expression-based clustering nor on the distribution and numbers of nuclei per cell sub-type across sleep conditions (Fig. 2 C). However, the SD condition was associated with a significantly increased number of transcripts/nucleus across sleep condition, matched for cell subtype (2way ANOVA, UMI/Nucleus(subType) x Condition(CSvsSD), F (16, 96) = 3.004, p=0.0004; Wilcoxon matched-pairs, p<= 0.0001; Table 2; Fig. 2 D). To avoid any technical biases, the CS and SD biological samples were treated in an identical manner, with sacrifice alternating between CS and SD samples, carried out by the same person, brain harvesting and frontal cortex isolation all carried out by Dr. Richa Pandey (blinded as to sample identity), RIN values were indistinguishable and finally, library preparation and sequencing were all completed in a single batch. Thus, the results may be indicative of a generalized increase in transcript number per cell sub-type.

Revised version: Approximation and Projection (UMAP) and annotated using multi-modal characteristics from the Brain Initiative Cell Census(Network, 2021) (Fig. 2 A,B). The sleep condition had no significant effect on the expression-based clustering nor on the distribution and numbers of nuclei per cell sub-type across sleep conditions (Fig. 2 C). However, the SD condition was associated with a significantly increased number of transcripts/nucleus across sleep condition, matched for cell subtype (2way ANOVA, UMI/Nucleus(subType) x Condition(CSvsSD), F (16, 96) = 3.004, p=0.0004; Wilcoxon matched-pairs, p<= 0.0001; Table S2; Fig. 2 D). To avoid any technical biases, the CS and SD biological samples were treated in an identical manner, with sacrifice alternating between CS and SD samples, carried out by the same person, brain harvesting and frontal cortex isolation all carried out by Dr. Richa Pandey (blinded as to sample identity), RIN values were indistinguishable and finally, library preparation and sequencing were all completed in a single batch. Thus, the results may be indicative of a generalized increase in transcript number per cell sub-type.

Earlier version: snRNAseq data shows cell type and sub-type based on gene expression patterns are unaffected by sleep need. A) UMAP projection of cell-type gene expression pattern following 6 hours ad lib sleep (CS) at ZT=6hour. B) As in “A” except after 6hours sleep deprivation (SD), ZT=6hour. C) The distribution of cell numbers across subtypes is unaffected by sleep need. D) The median number of UMIs/cell is significantly increased by sleep need across all cell subtypes (see_Table 2 and text for statistics).

Revised version: snRNAseq data shows cell type and sub-type based on gene expression patterns are unaffected by sleep need. A) UMAP projection of cell-type gene expression pattern following 6 hours ad lib sleep (CS) at ZT=6hour. B) As in “A” except after 6hours sleep deprivation (SD), ZT=6hour. C) The distribution of cell numbers across subtypes is unaffected by sleep need. D) The median number of UMIs/cell is significantly increased by sleep need across all cell subtypes (see_Table S2 and text for statistics).

Earlier version: We next determined the differential gene expression (DGE) of transcripts between conditions of CS and SD (see methods; Fig. 3 A). In response to 6 hours of SD from ZT= 0 to 6, both increased and decreased gene expression was observed in all cell subtypes (Table 2, 3). Predictably, there is a trend for increased DEGs in cell types expressing a greater number of genes (Fig. 3 A). The greatest proportion of sleep DEGs is ∼69%, expressed by ExIT cell types, while the next largest proportion is ∼17% for IN cell types (Fig. 3 B). To determine whether the increased numbers of genes expressed in the ExIT subtypes accounts for this cell type’s high proportion of DEGs, we examined each cell type’s probability of DEG expression given the number of genes expressed in the cell type, its Bayesian probability of expression.

Revised version: We next determined the differential gene expression (DGE) of transcripts between conditions of CS and SD (see methods; Fig. 3 A). In response to 6 hours of SD from ZT= 0 to 6, both increased and decreased gene expression was observed in all cell subtypes (Table S2, S3). Predictably, there is a trend for increased DEGs in cell types expressing a greater number of genes (Fig. 3 A). The greatest proportion of sleep DEGs is ∼69%, expressed by ExIT cell types, while the next largest proportion is ∼17% for IN cell types (Fig. 3 B). To determine whether the increased numbers of genes expressed in the ExIT subtypes accounts for this cell type’s high proportion of DEGs, we examined each cell type’s probability of DEG expression given the number of genes expressed in the cell type, its Bayesian probability of expression.

Cell type specific differential gene expression in response to 6 hours of SD. A) An XY bar plot of cell type specific DEGs (both up and down regulated) organized by cell type specific number of expressed genes (X axis) shows the greatest number of DEGs are found in excitatory pyramidal neurons that project within the telencephalon (ExIT). B) A pie chart of the distribution of DEGs amongst different cell types illustrates that the majority of sleep DEGs are expressed by ExIT neurons. C) An analysis of the cell type specific DEG occurrence shows the greatest probability of significant sleep loss gene response is found in ExIT cells by more than 3-fold compared to all other cell types.

The Bayesian probability of sleep loss DEGs, for cell subtype, is more than three-fold greater for excitatory intratelencephalically projecting neurons (ExIT subtype, pyramidal cells (Network, 2021)) compared to any other class of cells (Fig. 3 C). These observations suggest, with respect to transcriptomic mediated changes in the frontal cortex, the ExIT class neurons comprise the major target of sleep function and extends earlier observations that frontal glutamatergic neurons are the primary targets of the SD transcriptomic response (Bjorness et al., 2020; Kim et al., 2022).

We determined the functional relevance of cell-type specific DEGs with respect to both disease risk and biological properties. We found class-selective enrichment of both autism spectrum disorder genes (ASD; Fig. 4 A) and synaptic shaping components (Koopmans et al., 2019) (Fig. 4 A.).

DEG enrichment of cell types in response to SD by autism risk genes, synaptic shaping component genes and DEGs from MEF2c loss of function or constitutive HD4 repression of MEF2c. A) Heat map for cell type DEG enrichment by ASD risk genes, Synaptic Shaping Component genes, MEF2c-cKO DEGs and HD4cn DEGs. B) Model for the control of sleep DEGs by HD4 repression of MEF2c and by pMEF2c during low sleep need and, during high sleep need, de-repression of MEF2c by sequestration of pHD4 to the cytoplasm and dephosphorylation of MEF2c, facilitating expression of sleep genes. C) Model of downstream control by MEF2c and HD4 to facilitate SSC gene expression as sleep need increases. The SSC gene expression after waking, can then decrease AMPA/NMDA ratio and increase silent synapses during sleep to bias glutamate synapses towards decreased strength but increased potential for LTP at the start of the active phase when sleep need is low. D) An illustration of the SWS-SWA response to chronic MEF2c repression or activation. Chronic activation of MEF2c facilitated transcription leads to decreased AMPAR-mediated synaptic strength mimicking increased Ado tone that inhibits cortical-thalamic, glutamate synaptic activity, decreasing activation of cholinergic arousal centers and increasing SWS-SWA. Chronic repression of MEF2c does the opposite, mimicking loss of ADORA1 function and decreasing SWS-SWA.

Enrichment of sleep responsive DEGs by SSC and ASD genes, is strongest for ExIT (L 2-3, 4-5, 5) pyramidal neurons (Chi square with Yates correction p<0.006). In particular, a number of these sleep need ExIT-DEGs encode proteins that can directly affect AMPA/NMDA ratio and silent synapse numbers, including Dcc (Horn et al., 2013), Dgkb (Kakefuda et al., 2016), Gpc6 (Sato et al., 2016), Grin3a (Perez-Otano et al., 2016), Kif17 (Setou et al., 2000; Iwata et al., 2020), Kirrl3 (Martin et al., 2015) and Ptprf (Sclip and Sudhof, 2020). Some sleep modulated SSCs are also ASD risk factors in frontal cortical EXIT pyramidal neurons, including Cdh13, Dcc, Glra2, Gpc6, Grik4, Itpr1, Kirrel3, Nr1d1, Pcdh15 and Ush2a. To a much lesser extent, Layer 6 corticothalamic pyramidal neurons and somatostatin, parvalbumin and Sncg classes of inhibitory neurons (Scala et al., 2021) also show this enrichment of their sleep DEGs (Figure 4 A). These observations are consistent with the ExIT pyramidal cell, glutamatergic, synaptic phenotype as a major functional target for sleep proteins and implicates sleep (dis-)function’s role in autism risk.

Removed in the revised version: The transcription factor, MEF2c is necessary for sleep loss gene expression (Bjorness et al., 2020). ExIT sleep loss genes are significantly (p<0.0001, Chi2 with Yates correction) enriched with MEF2c target genes (Harrington et al., 2016; Bjorness et al., 2020) (Figure 4 A.). Class II histone deacetylases 4/5 (HD4, HD5) are MEF2c binding partners and repressors (McKinsey et al., 2000), that can shuttle in and out of the nucleus and are trapped in the cytoplasm when phosphorylated by SIKinases (salt induced kinases (Kim et al., 2022; Zhou et al., 2022)), known to become active in conditions of high sleep need. As the mouse sleep phase progresses, HD4/5 is progressively de-phosphorylated in association with decreased sleep need and conversely as the mouse active phase progress, HD4/5 is progressively phosphorylated in association with increased sleep need (see Figure 4 in Zhou, et al. (Zhou et al., 2022)).

Earlier version: A virally-mediated expression of mutant form of HD4, that resists phosphorylation and is thus at constitutively high levels in the nucleus (HD4cn (Zhou et al., 2022)), is expected to mimic a loss of function of MEF2c through constitutive repression of MEF2c. The sleep loss genes observed in frontal cortical ExIT neurons are enriched for both conditional Mef2c (Bjorness et al., 2020) loss of function DEGs (p<0.0001, Chi2 with Yates correction) and for DEGs in response to AAV-mediated expression of HD4cn (Zhou et al., 2022) (p<0.0001, Chi2 with Yates correction), as illustrated in a sleep-transcriptome expression model (Figure 4 B.). Further, these findings are indicative of the role of HD4 as a repressor of MEF2c facilitated sleep gene expression when sleep need is low and its de-repression of MEF2c through phosphorylation induced sequestration into the cytoplasm when sleep need is high (Zhou et al., 2022). The resulting high-sleep-need DEGs as noted above, can mediate the observed recovery of glutamate synapse phenotype from a high AMPA/NMDA ratio with sparse silent synapses, to lowered AMPA/NMDA ratio and increased silent synapses (Fig. 4 C).

Revised version: The transcription factor, MEF2c is necessary for sleep loss gene expression (Bjorness et al., 2020). ExIT sleep loss genes are significantly (p<0.0001, Chi2 with Yates correction) enriched with MEF2c target genes (Harrington et al., 2016; Bjorness et al., 2020) (Figure 4 A.). Class II histone deacetylases 4/5 (HD4, HD5) are MEF2c binding partners and repressors (McKinsey et al., 2000), that can shuttle in and out of the nucleus and are trapped in the cytoplasm when phosphorylated by SIKinases (salt induced kinases (Kim et al., 2022; Zhou et al., 2022)), known to become active in conditions of high sleep need. As the mouse sleep phase progresses, HD4/5 is progressively de-phosphorylated in association with decreased sleep need and conversely as the mouse active phase progress, HD4/5 is progressively phosphorylated in association with increased sleep need (see Figure 4 in Zhou, et al. (Zhou et al., 2022)). A virally-mediated expression of mutant form of HD4, that resists phosphorylation and is thus at constitutively high levels in the nucleus (HD4cn (Zhou et al., 2022)), is expected to mimic a loss of function of MEF2c through constitutive repression of MEF2c. The sleep loss genes observed in frontal cortical ExIT neurons are enriched for both conditional Mef2c (Bjorness et al., 2020) loss of function DEGs (p<0.0001, Chi2 with Yates correction) and for DEGs in response to expression of HD4cn (Zhou et al., 2022) (p<0.0001, Chi2 with Yates correction), as illustrated in a sleep-transcriptome expression model (Figure 4 B.). Further, these findings are indicative of the role of HD4 as a repressor of MEF2c facilitated sleep gene expression when sleep need is low and its derepression of MEF2c through phosphorylation induced sequestration into the cytoplasm when sleep need is high (Zhou et al., 2022). The resulting high-sleep-need DEGs as noted above, can mediate the observed recovery of glutamate synapse phenotype from a high AMPA/NMDA ratio with sparse silent synapses, to lowered AMPA/NMDA ratio and increased silent synapses (Fig. 4 C).

Discussion

Prolonged waking is associated with increased glutamatergic synaptic strength (Vyazovskiy et al., 2008; Liu et al., 2010; Bjorness et al., 2020) and increased number of functional AMPAR synapses (Bjorness et al., 2020). Both effects are reversed by recovery sleep (Bjorness et al., 2020). The evidence presented here suggests this wake/sleep homeostasis also reflects a waking induced increase in functional AMPA/NMDA ratio of glutamate synapses of the frontal cortex together with loss of silent glutamatergic synapses. The latter suggests a reduced availability of glutamate synapses for conversion from silent to active, reflecting saturation of LTP. Accordingly, it is reasonable to summarize the implications of sleep-related glutamate synaptic phenotype oscillation as: 1) prolonged waking induces a generalized synaptic strengthening coupled with a negative bias against potentiating plasticity as available slots for AMPARs saturate, and 2) a recovery by sleep (Fig. 4 C).

Learning of motor tasks that can be robustly and quantitatively monitored, like bird song in juvenile zebra finches (Deregnaucourt et al., 2005; Kollmorgen et al., 2020), may reflect the sleep/wake oscillation of bias for LTP and synaptic strengthening. At the start of the active phase, facilitated motor learning improves song performance to a level that appears to saturate by the end of the day. After a night’s sleep, performance is degraded but shows enhanced learning, facilitated by decreased AMPA/NMDA ratio and increased availability of synapses that can be readily potentiated. Since there are inevitable, day to day, unrelated variances in the environment that should not be learned, mitigation of over-fitting a learned task may require an oscillatory, stairstep-like, learning process. In rats, performance during learning of a novel motor task can display these similar characteristics, especially at the start of learning a novel task when exploration of a multi-dimensional learning space is at a premium (Kim et al., 2023), consistent with an underlying sleep/wake oscillation of glutamate synaptic phenotype.

Earlier version: Our unbiased examination of the ontology of ExIT DEGs indicates a response to sleep loss to modify ExIT glutamate synapses involving SSC DEGs. Our electrophysiological observations, now show more particularly, that the gene-expression modifications of ExIT neurons are coupled to SD-responses of, 1) increased AMPA/NMDA ratio; 2) decreased functional silent synapses and; 3) the previously described, increase in synaptic strength and functional number (Fig. 1 and 4).

Revised version: Our unbiased examination of the ontology of ExIT DEGs indicates a response to sleep loss to modify ExIT glutamate synapses involves SSC DEGs. Our electrophysiological observations, now show more particularly, that the gene-expression modifications of ExIT neurons are coupled to SD-responses of, 1) increased AMPA/NMDA ratio; 2) decreased functional silent synapses and; 3) the previously described, increase in synaptic strength and functional number (Fig. 1 and 4).

The ontology of the ExIT SD-response transcriptome is also notable for enrichment by autism-related risk genes, many of which overlap with our curated category of SSCs. A similar overlap has been noted for social affiliative behaviors, synaptic adhesion molecules, other synaptic shaping molecules with autism/autism spectrum disorder (ASD) risk genes in the mammalian CNS (Taylor et al., 2020). The association of sleep disruption with autism is generally appreciated (Mazurek et al., 2019) but the focus was on loss of time spent asleep or disrupted sleep, rather than loss of sleep function. The overlapped SSC and ASD DEGs provide a molecular association between sleep’s functional role in motor cortex and ASD risk that captures ASD’s association with motor deficits observed in patients (Chukoskie et al., 2013) and in mouse autism models (Cording and Bateup, 2023).

The molecular mechanisms responsible for sleep/wake DEGs observed in ExIT neurons include two critical transcription factors, MEF2c (Bjorness et al., 2020) and HD4/5 (Kim et al., 2022; Zhou et al., 2022). All neuronal and non-neuronal frontal cortical transcriptomic changes are abolished by loss of function of Mef2c (Bjorness et al., 2020). The abolishment of the sleep transcriptome remained even after restriction of the Mef2c knockout to CamKII-expressing glutamatergic neurons of the forebrain. Importantly, MEF2c was necessary for both the SD-transcriptomic response and for SD-induced increase of glutamatergic synaptic strength and functional synaptic number as well as for ensuing sleep-mediated recovery (Bjorness et al., 2020). Thus, MEF2c has an essential role in facilitating the expression of sleep genes needed for functional recovery from loss of sleep.

The activity of either of these factors to alter the sleep transcriptome, is controlled by their phosphorylation state. When sleep need is high, MEF2c is de-phosphorylated (Bjorness et al., 2020) (from pMEF2c to MEF2c) and Salt Induced Kinases (SIKs) phosphorylate HD4 and HD5 (pHD4/5), sequestering them to the cytoplasm (Kim et al., 2022; Zhou et al., 2022). Both the phosphorylation of HD4/5s (Miska et al., 1999) and de-phosphorylation of MEF2c (Zhu and Gulick, 2004), de-represses MEF2c transcriptional activity. The de-repression of MEF2c and its dephosphorylation are essential for SD-induced differential expression (Bjorness et al., 2020). The target genes of MEF2c activation can be inferred from the cortical, differential transcriptome induced by MEF2c loss of function (Harrington et al., 2016). These same genes overlap with the differentially expressed genes observed in response to overexpression of constitutively nuclear HD4 (HD4cn). This is predictable as HD4cn will constitutively bind and repress MEF2c in the nucleus, thus mimicking a MEF2c loss of function. Furthermore, in association with low sleep need, nuclear, and thus transcriptionally active (as a repressor), HD4, is more abundant in a de-phosphorylated state, relative to transcriptionally inactive, phosphorylated cytoplasmic pHD4 (Kim et al., 2022; Zhou et al., 2022). NMDAR activation negatively controls HD4 transcriptional activity and nuclear localization (Sando et al., 2012). Accordingly, as NMDAR activity accumulates during the active period, an expected accumulation of pHD4 in the cytoplasm occurs to derepress MEF2c activity. Thus, HD4 may act to repress MEF2c transcriptional facilitation of sleep genes, during low sleep need and derepress MEF2c as sleep need builds during the active period. We observed DEGs from both MEF2c loss of function and HD4cn also sleep DEGs consistent with their interaction driven by sleep need (Figure 4 C).

Both the role of MEF2c to mediate sleep gene expression and nuclear HD4 to repress MEF2c activity, lead to an apparent paradox. When HD4 repression of MEF2c is chronically reduced or lost, sleep need, as indicated by slow wave activity during slow wave sleep (SWS-SWA), is increased (Kim et al., 2022). If MEF2c promotes sleep gene expression, then sleep need should be reduced not increased by its de-repression. Conversely, when constitutively, nuclear HD4 is overexpressed, to chronically repress MEF2c, sleep need, as indicated by SWS-SWA, is unexpectedly, reduced (Zhou et al., 2022). But, what is actually being indicated by SWS-SWA?

Earlier version: SWS-SWA has long been employed as a marker for sleep need or intensity, primarily because of its strong correlation with previous time spent awake (Franken et al., 2001; Borbely et al., 2016). Despite this marker’s long history, the mechanisms responsible for its correlation to sleep need are not well understood. A likely candidate is adenosine (Ado) since, it is increased by sleep-need (Porkka-Heiskanen et al., 1997) and is necessary for SWS-SWA rebound in response to SD (Bjorness et al., 2009). Ado is increased by glutamate synaptic activity acting on NMDARs in arousal centers. Ado activation of Ado A1 receptors (ADORA1) inhibits AMPAR synaptic activity (Brambilla et al., 2005; Greene et al., 2017). Adk, which encodes adenosine kinase in glial cells is the high affinity metabolizing enzyme of Ado. Glial knockout of Adk increases extracellular Ado and SWA in both wake and sleep (Bjorness et al., 2016). Rebound SWS-SWA in response to experimental SD increases Ado (Porkka-Heiskanen et al., 1997) and the rebound is blocked by a CamKII:Cre driven conditional knockout of Adora1 (Bjorness et al., 2016). Thus Ado tone, whether increased by prolonged waking activity or by glial knockdown of Adk, through its activation of ADORA1, inhibits glutamate synaptic activity, including in arousal centers, to promote SWS-SWA (Greene et al., 2017) (Figure 4 D).

Revised version: SWS-SWA has long been employed as a marker for sleep need or intensity, primarily because of its strong correlation with previous time spent awake (Franken et al., 2001; Borbely et al., 2016). Despite this marker’s long history, the mechanisms responsible for its correlation to sleep need are not well understood. A likely candidate is adenosine (Ado) since, it is increased by sleep-need and is necessary for SWS-SWA rebound in response to SD (Porkka-Heiskanen et al., 1997; Bjorness et al., 2009; Bjorness et al., 2016). Ado is increased by glutamate synaptic activity acting on NMDARs in arousal centers. Ado activation of Ado A1 receptors (ADORA1) inhibits AMPAR synaptic activity (Brambilla et al., 2005; Greene et al., 2017). Adk, which encodes adenosine kinase in glial cells is the high affinity metabolizing enzyme of Ado. Glial knockout of Adk increases extracellular Ado and SWA in both wake and sleep (Bjorness et al., 2016). Rebound SWS-SWA in response to experimental SD increases Ado (Porkka-Heiskanen et al., 1997) and the rebound is blocked by a CamKII:Cre driven conditional knockout of Adora1 (Bjorness et al., 2016). Thus Ado tone, whether increased by prolonged waking activity or by glial knockdown of Adk, through its activation of ADORA1, inhibits glutamate synaptic activity, including in arousal centers, to promote SWS-SWA (Greene et al., 2017) (Figure 4 D).

Earlier version: Sleep, itself, induces a down-regulation of glutamate synaptic strength, synaptic numbers, and AMPA/NMDA ratio. Chronic facilitation of MEF2c-dependent sleep genes due to de-repression of MEF2c by HD4/5’s loss of function should induce a chronic down-regulation of glutamatergic synaptic drive (Kim et al., 2022; Zhou et al., 2022). The sleep-gene induced, chronic glutamate synaptic down-regulation can mimic the effect of a tonic increase of extracellular Ado to increase SWA as occurred following glial loss of function of Adk (Bjorness et al., 2016). Furthermore, synaptic glutamate activity in cholinergic arousal centers (Rainnie et al., 1994; Brambilla et al., 2005) can be reduced, decreasing cholinergic mediated thalamo-cortical activation, and increasing SWS-SWA.

Revised version: Sleep, itself, induces a down-regulation of glutamate synaptic strength, synaptic numbers, and AMPA/NMDA ratio. Chronic facilitation of MEF2c-dependent sleep genes due to de-repression of MEF2c by HD4/5’s loss of function should induce a chronic down-regulation of glutamatergic synaptic strength and number (Bjorness et al., 2020; Kim et al., 2022; Zhou et al., 2022). The sleep-gene induced, chronic glutamate synaptic down-regulation can mimic the effect of a tonic increase of extracellular Ado to increase SWA as occurred following glial loss of function of Adk (Bjorness et al., 2016). Furthermore, synaptic glutamate activity in cholinergic arousal centers (Rainnie et al., 1994; Brambilla et al., 2005) can be reduced, decreasing cholinergic mediated thalamo-cortical activation, and increasing SWS-SWA.

Conversely, repression of MEF2c by de-phosphorylated HD4/5 (or by expression of phospho-dead HD4cn) reduces SWS-SWA (Zhou et al., 2022). In this case, glutamate synaptic activity is chronically increased (directly increasing C-T activation and indirectly driving cholinergic arousal center activity) by the direct effect of loss of sleep function, as evidenced by a loss of function of Mef2c in glutamate neurons that prevents sleep gene expression (Bjorness et al., 2020). This prevents the normal down-regulation of glutamate synaptic function (Bjorness et al., 2020) and can mimic the effect of a loss of Adora1 function on SWS-SWA (Bjorness et al., 2009). It predicts HD4cn’s repressive effect on sleep gene expression in response to SD, consistent with the enrichment of SD genes by both HD4cn and Mef2c-cKO DEGs, that we observed (Fig. 4 A).

Added in the revised version: Generally, facilitation or depression of SWA correlates with up or down scaling effects on cortical glutamate neurotransmission, respectively, even in the absence of a direct effects on sleep need (Figure 4D). Many reagents that reduce the excitability of glutamate pyramidal cells by various mechanisms, including anesthetics like isoflurane, barbiturates or benzodiazepines in addition to those activating ADORA1, increase SWA. Those agents that increase excitability of cortical glutamate pyramidal neurons, correlate with decreased SWA. Accordingly, SWA can be a marker for glutamate synaptic strength and number, modulated under normal conditions, by sleep function.

In summary, at the end of a long episode of waking, when sleep need is high, glutamate synapses of ExIT cells in the frontal cortex are electro-physiologically strengthened, increased in functional number, and show increased AMPA/NMDA ratio, together with decreased availability of silent synapses. The silent synapse conversion by LTP to active synapses is thus limited, creating a bias for increased strength at the expense of potentiating plasticity. In the ExIT neuronal nuclei, the transcription factor, MEF2c is de-repressed by dephosphorylation and sequestration of its co-repressors, the classII HDACs, pHD4/5, to the cytoplasm. This facilitates transcription of sleep genes, including those encoding SSCs controlling glutamatergic synaptic phenotype and ASD risk genes. Recovery sleep recovers the functional phenotype. These observations suggest a daily oscillation from lowered glutamatergic synaptic strength and increased bias for potentiating plasticity at the start of the active phase when sleep need is low, to increased synaptic strength and saturated plasticity at the end of the active phase, when sleep need is high. In conclusion, we have provided electrophysiological evidence for a wake/sleep oscillation of glutamate synapse phenotype that can mediate a glutamatergic strength/plasticity oscillating bias. We also show that this wake/sleep oscillation is likely to occur in glutamate synapses of ExIT pyramidal neurons and be mediated by a select set of synaptic shaping component genes, a significant number of which are also autism risk genes, whose expression is controlled by MEF2C and HD4 transcription factors. Finally, this study implicates a framework within which optimal cortical-dependent motor training can occur in a recursive incremental manner, facilitated by wake/sleep glutamate synapse phenotypical oscillation.

Materials and methods

Electrophysiology

All electrophysiological animal experiments were approved by the Animal Experimental Committee at the University of Tsukuba. Animals were housed on a 12:12-h light/dark cycle (with the light automatically turned on at 9:00) at a stable temperature (24.5 ± 0.5 °C), with free access to food and water. C57/Bl6 male mice aged 8 to 12 weeks were used.

Slice preparation and recording

Earlier version: Mice were fully anesthetized at ZT 6 with isoflurane, and their brains were promptly removed (<1 min) and placed in ice-cold artificial cerebrospinal fluid (ACSF), composed of 124 mM NaCl, 26 mM NaHCO3, 3 mM KCl, 2 mM CaCl, 1 mM MgSO4, 1.25 mM KH2PO4, 10 mM glucose, 300 to 310 mOsm, equilibrated with 95% O2 and 5% CO2.

Revised version: Mice were fully anesthetized at ZT 6 with isoflurane, and their brains were promptly removed (<1 min) and placed in ice-cold artificial cerebrospinal fluid (ACSF), composed of 124 mM NaCl, 26 mM NaHCO3, 3mM KCl, 2 mM CaCl, 1 mM MgSO4, 1.25 mM KH2PO4, 10 mM glucose, 300 to 310 mOsm, equilibrated with 95% O2 and 5% CO2.

Earlier version: Brains were mounted on the stage of a vibrating microtome (Leica VT1200 S) with cyanoacrylate glue. Coronal sections (300 μm thick) were obtained and allowed to recover in ACSF at room temperature for at least 1 h. Slices were submerged in a perfusion chamber placed under an upright microscope (BX51WI; Olympus) fitted with a custom LED IR illumination and Nomarski interference contrast. Slices were superfused with ACSF, at a rate of 2 mL/min with Picrotoxin (100 micromolar) added to block GABAA receptor-mediated responses. Neurons in layers 2/3 of primary motor cortex with a pyramidal cell morphology were patched under visual control. The intracellular solution contained (in mM) 130 Cs gluconate, 1 mM EGTA, 10 mM Hepes, 4 mM MgATP, 0.3 mM NaGTP, and 5 mM NaCl. The pH was adjusted to 7.33 with CsOH and osmolarity to 285 to 300 mOsm. Borosilicate patch pipettes were pulled to an open-tip resistance of 2 to 4 MΩ. Signals were amplified and filtered at 10 kHz (Axopatch 700A; Molecular Devices) and then digitized at 20 kHz using customized routines in commercial software (IGOR Pro; WaveMetrics). Evoked synaptic responses were produced through extracellular stimulation (0.1 ms) every 10 s with monopolar glass pipette electrodes filled with ACSF and placed in the vicinity (200-400 micrometers) lateral to the recorded neuron in layer 2/3.

Revised version: Brains were mounted on the stage of a vibrating microtome (Leica VT1200 S) with cyanoacrylate glue. Coronal sections (300 μm thick) were obtained and allowed to recover in ACSF at room temperature for at least 1 h. Slices were submerged in a perfusion chamber placed under an upright microscope (BX51WI; Olympus) fitted with a custom LED IR illumination and Nomarski interference contrast. Slices were superfused with ACSF, at a rate of 2 mL/min with Picrotoxin (100 micromolar) added to block GABAA receptor-mediated responses. Neurons in layers 2/3 of primary motor cortex with a pyramidal cell morphology were patched under visual control. The intracellular solution contained (in mM) 130 Cs gluconate, 1 mM EGTA, 10 mM Hepes, 4 mM MgATP, 0.3 mM NaGTP, and 5 mM NaCl. The pH was adjusted to 7.33 with CsOH and osmolarity to 285 to 300 mOsm. Borosilicate patch pipettes were pulled to an open-tip resistance of 2 to 4 MΩ. Signals were amplified and filtered at 10 kHz (Axopatch 700A; Molecular Devices) and then digitized at 20 kHz using customized routines in commercial software (IGORPro; WaveMetrics). Series resistance values for the recording pipette ranged between 8 and 15 MOhm and experiments with changes larger than 25% were not used for further analyses.

Added in the revised version: Evoked synaptic responses were produced through extracellular stimulation (0.1 ms) every 10 s with monopolar glass pipette electrodes filled with ACSF and placed in the vicinity (200-400 micrometers) lateral to the recorded neuron in layer 2/3. For the AMPA/NMDA amplitude ratio experiments stimuli amplitudes were 8.7 uA +/- 3.5 uA (SD) and for the failure rate experiments, 1.4 uA +/-0.9 uA (SD).

AMPA/NMDA EPSC analyses

Added in the revised version: AMPA responses were recorded at a holding potential of -90 mV and mixed AMPA-NMDA responses at a holding potential of +50 mV. The -90 mV holding potential was chosen according to precedent (Myme et al., 2003). The hyperpolarized holding potential increases driving force and permits lower stimulus strength for the same response size – reducing the likelihood for polysynaptic responses. Experiments with multiple response peaks at -90 mV were not included in the analysis. The -90 mV holding potential also increases NMDA receptor Mg block - resulting in a minimally contaminated AMPA response.

Earlier version: AMPA responses were recorded at a holding potential of -90 mV and mixed AMPA-NMDA responses at a holding potential of +50 mV. The AMPA/NMDA ratio was determined according to an established protocol(Myme et al., 2003); briefly, at least 10 traces were averaged at -90 mV and +50 mV, the AMPA response was calculated as the peak inward current at -90 mV and the NMDA current was calculated at +50 mV in a 5 ms window 40 ms after the AMPA peak.

Revised version: The AMPA/NMDA ratio was determined according to an established protocol (Myme et al., 2003); briefly, at least 10 traces were averaged at -90 mV and +50 mV, the AMPA response was calculated as the peak inward current at -90 mV and the NMDA current was calculated at +50mV in a 5 ms window 40 ms after the AMPA peak.

Earlier version: Failure rates for AMPA and AMPA+NMDA EPSCs: Minimal stimulation was achieved by lowering the stimulation strength until clear successes and failures could be observed at -90 mV holding potential and stimulation strength was then kept constant for 5-10 min to allow the system to stabilize before failure ratios were determined. Relative numbers of silent and total (silent + active) synapse in an L2-3 pyramidal neuron evoked in a slice by minimal presynaptic stimulation are correlated with failure rates (FR= failures/number of stimuli). FR-90 = (1-prelease)As, and FR+50 = (1-prelease)As+Ss, where As=number active synapses (at -90mV) and Ss= number of silent synapses (only active at +50mV). Since prelease for each synapse is the same for -90mV and +50mV (i.e., assuming the same probability of failure at both potentials with each stimulus) then: Thus, Since (FR+50/FR-90)= (1-prelease)As+Ss / (1-prelease)As, as Ss-->0, (FR+50/FR-90)--> 1, the maximum ratio possible if prelease is constant. Our uncorrected data show (FR+50/FR-90)=∼1.3 in SD, which suggests that FR+50 is overestimated by at least a factor of 1.3 (likely, due to decreased input resistance at Vh=+50mV). Accordingly, we adjusted FR+50 by this factor = 1.3, for all sleep conditions.

Revised version: Failure rates for AMPA and AMPA+NMDA EPSCs: Minimal stimulation was achieved by lowering the stimulation strength until clear successes and failures could be observed at -90 mV holding potential and stimulation strength was then kept constant for 5-10 min to allow the system to stabilize before failure ratios were determined. Relative numbers of silent and total (silent + active) synapse in an L2-3 pyramidal neuron evoked in a slice by minimal presynaptic stimulation are correlated with failure rates (FR=failures/number of stimuli). FR-90 = (1-prelease)As, and FR+50 = (1-prelease)As+Ss, where As=number active synapses (at -90mV) and Ss= number of silent synapses (only active at +50mV). Since prelease for each synapse is the same for -90mV and +50mV (i.e., assuming the same probability of failure at both potentials with each stimulus) then:

Added in the revised version: ln(FR+50) = ln (1-prelease)As+Ss

Added in the revised version: = (As+Ss)*ln(1-prelease) ln(FR-90)= ln(1-prelease)As

Added in the revised version: =As*ln(pfailure) Thus,

Added in the revised version: ln(FR+50)/ln(FR-90) = (As+Ss)/As

Added in the revised version: =1 + Ss/As.

Added in the revised version: Since (FR+50/FR-90)= (1-prelease)As+Ss / (1-prelease)As, as Ss-->0, (FR+50/FR-90)--> 1, the maximum ratio possible if prelease is constant. Our uncorrected data show (FR+50/FR-90)=∼1.3 in SD, which suggests that FR+50 is overestimated by at least a factor of 1.3 (likely, due to decreased input resistance at Vh=+50mV). Accordingly, we adjusted FR+50 by this factor = 1.3, for all sleep conditions.

Sleep deprivation

Earlier version: All mice were accommodated for 1-2 weeks to a cage and treadmill, consisting of a bottomless plexiglass enclosure suspended over the treadmill with ad lib access to food and water. For sleep deprivation (SD) the slow-moving treadmill (0.1KM/hour) was engaged from ZT 0-6 hours and for recovery sleep (RS) from ZT 0-4. There was no treadmill engagement for RS from ZT 4-6 hours and for control sleep (CS), no treadmill engagement for the entire ZT 0-6 hours period. All animals were sacrificed at ZT=6 hours and tissue rapidly prepared for either electrophysiological (exVivo brain slices) or transcriptomic analysis as described below.

Revised version: All mice were accommodated for 1-2 weeks to a cage and treadmill, consisting of a bottomless plexiglass enclosure suspended over the treadmill with ad lib access to food and water. During this time, all animals were exposed to treadmill movement at ZT=2 and ZT=14 hours, for 30 minutes. For sleep deprivation (SD) the slow-moving treadmill (SMTM; 0.1KM/hour; see “video 1”) was engaged from ZT 0-6 hours and for recovery sleep (RS) from ZT 0-4. There was no treadmill engagement for RS from ZT 4-6 hours and for control sleep (CS), no treadmill engagement for the entire ZT 0-6 hours period. All animals were sacrificed at ZT=6 hours and tissue rapidly prepared for either electrophysiological (exVivo brain slices) or transcriptomic analysis as described below.

Single nucleus Transcriptome analysis

snRNA-seq library preparation

Nuclei for snRNA-seq were isolated from male mice, 8-10 weeks old. Briefly, the tissue was homogenized using a glass Dounce homogenizer in 2 ml of ice-cold lysis buffer (10 mM Tris-HCl, 10 mM NaCl, 3 mM MgCl2 and 0.1% Nonidet P40 Substitute) and was incubated on ice for 5 min. Nuclei were centrifuged at 500g for 5 min at 4 °C, washed with 4 ml ice-cold lysis buffer and incubated on ice for 5 min. Nuclei were centrifuged at 500g for 5 min at 4 °C. After centrifugation, the nuclei were resuspended in 500 μl of nucleus suspension buffer (NSB) containing 1× PBS, 1% BSA (no. AM2618, Thermo Fisher Scientific) and 0.2 U μl−1 RNAse inhibitor (no. AM2694, Thermo Fisher Scientific). The nucleus suspension was filtered through a 70-μm Flowmi cell strainer (no. H13680-0070, Bel-Art). Debris was removed with density gradient centrifugation using Nuclei PURE 2 M sucrose cushion solution and Nuclei PURE sucrose cushion buffer from the Nuclei PURE prep isolation kit (no. NUC201-1KT, Sigma Aldrich). Nuclei PURE 2 M sucrose cushion solution and Nuclei PURE sucrose cushion buffer were first mixed in a 9:1 ratio. A 500 μl volume of the resulting sucrose solution was added to a 2-ml Eppendorf tube. A 900 μl volume of the sucrose buffer was added to 500 μl of isolated nuclei in NSB. A 1,400 μl volume of nucleus suspension was layered to the top of the sucrose buffer. This gradient was centrifuged at 13,000g for 45 min at 4 °C. The pellet of nuclei was resuspended, washed once in NSB and filtered through a 70-μm Flowmi cell strainer (no. H13680-0070, Bel-Art). The concentration of nuclei was determined using 0.4% trypan blue (no. 15250061, Thermo Fisher Scientific), and was adjusted to a final concentration of 1,000 nuclei per microlitre with NSB.

Removed in the revised version: Droplet-based snRNA-seq libraries were prepared using Chromium Single Cell 3′ v3.1 (1000121, 10x Genomics) according to the manufacturer’s protocol(Zheng et al., 2017). Libraries were sequenced using an Illumina NovaSeq 6000.

Removed in the revised version: snRNA-seq preprocessing and annotationRaw sequencing data was obtained as BCL files from the McDermott sequencing core at UT Southewestern. BCL files were demultiplexed using cellranger mkfastq (10X Genomics CellRanger suite v3.1.0). Resulting FASTQ files were then accessed for quality using FASTQC (v0.11.5). A reference mouse genome-annotation index was built for mouse genome (GRCm38p6) and Gencode annotation (vM17) using cellranger mkref (10X Genomics CellRanger suite v3.1.0). Quality passed FASTQ files were further aligned to reference mouse genome-annotation index and raw count matrices were generated using cellranger count (10X Genomics CellRanger suite v3.1.0).

Removed in the revised version: To remove ambient RNA contamination, we used CellBender’s remove-background (https://github.com/broadinstitute/CellBender) on the raw count matrix per sample. We note that without ambient RNA removal, glial cells were shown to be conspicuously contaminated with neuronal ambient RNAs(Caglayan et al., 2022). Also, potential doublets were discarded using DoubletFinder (https://github.com/chris-mcginnis-ucsf/DoubletFinder) tool.

Earlier version: Ambient RNA and doublets cleaned data were used for downstream analysis. For each sample, nuclei with less than 20,000 UMIs and percentage of reads mapping to mitochondrial genes of less then 0.5 were retained. Individual samples per sleep condition were first integrated using Seurat’s (v3) integration approach (IntegrateData) and integrated datasets were further clustered (ScaleData, RunPCA, FindClusters) as described in Seurat’s (v3) integration vignette (https://satijalab.org/seurat/archive/v3.0/integration). Clustered data was then visualized (RunUMAP) using Uniform manifold approximation and projection embeddings (UMAP). Cluster-specific gene markers were identified (FindMarkers) and significant marker genes were enriched using Fisher exact test for cell-type markers defined in Brain Initiative Cell Census to annotate the cell-types. Cell-types were also confirmed by expression of canonical marker genes.

Revised version: Droplet-based snRNA-seq libraries were prepared using Chromium Single Cell 3’ v3.1 (1000121, 10x Genomics) according to the manufacturer’s protocol(Zheng et al., 2017). Libraries were sequenced using an Illumina NovaSeq 6000. snRNA-seq preprocessing and annotationRaw sequencing data was obtained as BCL files from the McDermott sequencing core at UT Southewestern. BCL files were demultiplexed using cellranger mkfastq (10X Genomics CellRanger suite v3.1.0). Resulting FASTQ files were then accessed for quality using FASTQC (v0.11.5). A reference mouse genome-annotation index was built for mouse genome (GRCm38p6) and Gencode annotation (vM17) using cellranger mkref (10X Genomics CellRanger suite v3.1.0). Quality passed FASTQ files were further aligned to reference mouse genome-annotation index and raw count matrices were generated using cellranger count (10X Genomics CellRanger suite v3.1.0). To remove ambient RNA contamination, we used CellBender’s remove-background (https://github.com/broadinstitute/CellBender) on the raw count matrix per sample. We note that without ambient RNA removal, glial cells were shown to be conspicuously contaminated with neuronal ambient RNAs(Caglayan et al., 2022). Also, potential doublets were discarded using DoubletFinder (https://github.com/chris-mcginnis-ucsf/DoubletFinder) tool. Ambient RNA and doublets cleaned data were used for downstream analysis. For each sample, nuclei with less than 20,000 UMIs and percentage of reads mapping to mitochondrial genes of less then 0.5 were retained. Individual samples per sleep condition were first integrated using Seurat’s (v3) integration approach (IntegrateData) and integrated datasets were further clustered (ScaleData, RunPCA, FindClusters) as described in Seurat’s (v3) integration vignette (https://satijalab.org/seurat/archive/v3.0/integration). Clustered data was then visualized (RunUMAP) using Uniform manifold approximation and projection embeddings (UMAP). Cluster-specific gene markers were identified (FindMarkers) and significant marker genes were enriched using Fisher exact test for cell-type markers defined in Brain Initiative Cell Census to annotate the cell-types. Cell-types were also confirmed by expression of canonical marker genes.

Differential gene expression tests were performed using edgeR-based pseudobulk approach across sleep conditions per cell-type. Significant differentially expressed genes (DEGs) were identified using absolute log2 fold change >= 0.1375 (10%) and a false discovery rate (FDR) <= 0.05. Further significant DEGs were enriched across defined gene classes (such as SSC genes, ASD genes, Mef2c-cKO genes and HD4cn genes) using R package SuperExactTest to functionally annotate the genes.

Statistics

The exact sample size (n) for each experimental group/condition, is provided as a discrete number and unit of measurement. Statistics were taken from distinct samples and not measured repeatedly, with the exception of the matched sampled evoked EPSCs recorded at Vh=-90mV and at Vh=+50mV from the same neuron as part of the analyses of AMPAR and NMDAR EPSC amplitudes and failure rates. Statistical details are provided in appropriate tables (Tables 1-3) in tabs labelled “Stats ….”, in the “Supplementary Information-Data/Tables”.

Removed in the revised version: Supporting information